You should *always* be able to "Kernel -> Restart and run all" the instructor versions of notebooks.  And should do this often.

```python
import os
import pandas
```

```python
if os.path.exists('/coursedata'):
    COURSEDATA = '/coursedata'
else:
    COURSEDATA = '.'
```

```python
# I copied /usr/share/dict/words to from my Linux workstation to the data directory
# using the workstation mount (renaming it to words.txt)
```

```python
# Read in the data
df = pandas.read_csv(os.path.join(COURSEDATA, 'words.txt'), names=["words"], header=None)
```

```python
df.head()
```

Output:
```
words
0      A
1    A's
2   AA's
3   AB's
4  ABM's
```

<!-- grade_id: cell-64c99be11824d217, points: 0, grade: false, solution: true -->
```python
# This is a student answer cell, so it must be marked "Autograded answer"

# Calculate the mean of all the length of all words
def mean_of_column(column):
    ### BEGIN SOLUTION
    mean = column.str.len().mean()
    ### END SOLUTION
    return mean
```

```python
# This cell is read-only, because why not?

# Let's make and test some testdata
df_test = pandas.DataFrame({"testwords": ["dog", "cat", "camel"]})
```

```python
mean_of_column(df_test['testwords'])
```

Output:
```
3.6666666666666665
```

```python
assert abs(mean_of_column(df_test['testwords']) - 3.66666 ) < .0001
```

```python
# This is a normal cell, that students can edit.  We hope they play around
# with things here and learn something.
mean_of_column(df['words'])
```

Output:
```
8.4643743067459916
```

<!-- grade_id: cell-b5655049d503321c, points: 0, grade: true, solution: false -->
```python
# This is marked "autograder tests", so can not be edited
assert abs(mean_of_column(df['words']) - 8.46437) < .0001
### BEGIN HIDDEN TESTS
df_test2 = pandas.DataFrame({"testwords": ["dog", "bird", "camel"]})
assert mean_of_column(df_test2['testwords']) == 4
### END HIDDEN TESTS
```

# Reflection
Please describe, in your own words, why the mean of the entire words file:

(this cell is "manually graded task" - markdown)

Please answer here

# Extension
Write a function to test standard deviation of all words:

<!-- grade_id: cell-b6a45d75ebb21b41, points: 0, grade: true, solution: true -->
```python
# "manually graded answer" - code

# This entire cell is replaced in the student version
```